A kidney-hypothalamus axis promotes compensatory glucose production in response to glycosuria

Author Response

The following is the authors’ response to the current reviews.

Response to Reviewer 1:

We agree with the reviewer’s overall assessment of this manuscript.

Because multiple secreted proteins are changed between the control and experimental groups, some of them could be causal and others corelative in the context of enhancing compensatory glucose production in response to elevated glycosuria. Through future studies we will determine the causal factors that trigger the increase in glucose production.

Yes, we will correct the typographical errors in a revised version of this manuscript.

Response to Reviewer 2:

We agree with reviewer on their comment about potential sex differences we may have missed in this study. Therefore, we will include this limitation in discussion section of a revised manuscript.

The reviewer’s statement ‘The methods of that publication indicate that all experiments were completed within 14 days of inducing the Glut2 knockout’ is incorrect. In the referred publication, we had explicitly mentioned in methods that ‘All of the experiments, except those using a diet-induced obesity mouse model or noted otherwise, were completed within 14 days of inducing the Glut2 deficiency.’ Please see figures 5h-l and 6 in that previous publication, which demonstrate that all the experiments were not completed within 14 days of inducing renal Glut2 deficiency. Per the reviewer’s advice, in the present manuscript we will include the timeline of the experiments (which in some cases is 4 months beyond inducing glycosuria) with all the figure legends. In addition, for a separate project (which is unpublished) we have measured glycosuria up to 1 year after inducing renal Glut2 deficiency. Therefore, the glycosuria observed in the renal Glut2 KO mice is not temporary.

In our previous response to the reviewer, we had already mentioned which control group was used in this study. Please see our response to the second reviewer’s point 3. As mentioned to the reviewer, we had used Glut2-loxp/loxp mice as the control group, which is also described multiple times in the figure legends of our previous paper that reported the phenotype of renal Glut2 KO mice and is cited in this manuscript so we don’t have to repeat the same information. Per the reviewer’s advice, we will also include the information in a revised version of this manuscript.

We request the reviewer to look at figure 1, showing an increase in glucose production in renal Glut2 KO mice and figure 3, which demonstrates that an afferent renal denervation reduces blood glucose levels by 50%. The afferent renal denervation (ablation of afferent renal nerves) does reduce blood glucose levels in renal Glut2 KO mice. Therefore, the use of the word ‘promote’ in the title is accurate and appropriate to reflect the role of the afferent renal nerves in contributing to about 50% increase in blood glucose levels in renal Glut2 KO mice. Regarding the reviewer's comment on changes in Crh gene expression, please look at figure 3. Ablation of renal afferent nerves decreases hypothalamic Crh gene expression and other mediators of the HPA axis by 50%. Therefore, the afferent renal nerves do contribute to regulating blood glucose levels, at least in part, by the HPA axis (which is widely known to change blood glucose levels). The use of words such as ‘required’ or ‘necessary’ in the title may have indicated causal role or could have been misleading here; therefore we have purposely used ‘promote’ in the title to accurately reflect the findings of this study.

Because we observed an increase in hepatic glucose production in renal Glut2 KO mice (Fig. 1) - which was reduced by 50% after selective afferent renal denervation (Fig. 3) - in the graphical abstract we are suggesting a neural connection between the kidney-brain-liver or an endocrine factor(s) to account for these changes in blood glucose levels as also described in the discussion section. We can include a question mark ‘?’ in the graphical abstract to show that further studies are need to validate these proposed mechanisms; however, we cannot just remove the arrow as advised by the reviewer.

Per the reviewer’s advice, in the methods we will include the dilutions used for each assay.

The following is the authors’ response to the original reviews.

Reviewer #1 (Recommendations For The Authors):

It would be helpful to the reader to specify in Figure 1a-c whether data were directly measured or calculated.

We have now clarified this in method section of the revised manuscript. The glucose production was directly measured and then fractional contribution of the tissues was calculated from the former data. We have also included a reference research paper to further clarify the method.

The methods section would be strengthened by clarifying the order in which experiments were performed, the age of the mice at each time point, and whether different cohorts were used for different techniques.

We have included additional details in the method section with proper citations. For in-depth protocols we have cited our previous publications.

It would be helpful to explain or provide a reference for how the post-mortem background activity measurement was performed.

We have included this explanation in the revised manuscript.

Similarly, details regarding the collection of blood for ACTH and corticosterone measurement are needed for the reader to evaluate whether the results are confounded by stress at the time of collection.

We have added these details in the method section.

I recommend stating, if accurate, that you used mixed-sex groups because your previous study found no sex differences in the phenotype of renal Glut2 KO mice.

Yes, we have included these details in the revised manuscript.

Sentence 239 is difficult to follow. Also, line 287 contains a contraction.

We have revised the sentence per the reviewer’s advice.

A graphical abstract would be helpful, bearing in mind conclusive vs suggestive findings.

Yes, we have included the graphical abstract with the revised manuscript.

Reviewer #2 (Recommendations For The Authors):

Minor Comments to the Authors

(1) The Methods also need to specify more of the critical details of the ELISAs, including the dilution factors used, and whether the values reported are dilution-corrected. Also, there is no description of how insulin was measured.

We have included these details in the method section. The assay dilutions were performed per manufacturers’ instructions.

(2) The Methods do not sufficiently describe how Crh mRNA was quantified in the hypothalamus. Presumably, they examined only the paraventricular nucleus? How many sections were used for in situ hybridization? How were the brains processed? What thickness of section was used? When were the brains collected?

We have included these details in the method section and cited our previous publications for in-depth protocols. Some of the information is also available in the figure legends.

(3) The number of mice that were used for plasma proteomics is not indicated.

The number of mice is indicated using individual symbols or points presented on the bar graphs.

Changes from the earlier version to the revised version

Abstract

Abstract

The kidneys facilitate energy conservation through reabsorption of nutrients including glucose. Almost all of the filtered blood glucose is reabsorbed by the kidneys. Loss of glucose in urine (glycosuria) is offset by an increase in endogenous glucose production to maintain normal energy supply in the body. How the body senses this glucose loss and consequently enhances glucose production is unclear. Using renal Glut2 knockout mice, we demonstrate that elevated glycosuria activates the hypothalamic-pituitary-adrenal axis, which in turn drives endogenous glucose production. This phenotype was attenuated by selective afferent renal denervation, indicating the involvement of the afferent nerves in promoting the compensatory increase in glucose production. In addition, through plasma proteomics analyses we observed that acute phase proteins - which are usually involved in body’s defense mechanisms against a threat – were the top candidates which were either upregulated or downregulated in renal Glut2 KO mice. Overall, afferent renal nerves contribute to promoting endogenous glucose production in response to elevated glycosuria and loss of glucose in urine is sensed as a biological threat in mice. These findings may be useful in improving efficiency of drugs like SGLT2 inhibitors that are intended to treat hyperglycemia by enhancing glycosuria, but are met with a compensatory increase in endogenous glucose production.

Introduction

The kidneys help conserve energy by reabsorbing nutrients and preventing their loss in urine. Almost all of the filtered glucose is reabsorbed by the kidneys. The renal threshold for glucose reabsorption in humans is 180 mg/dl and in rodents is about 400 mg/dl[1]. In conditions such as overt hyperglycemia (beyond the renal glucose threshold), genetic loss of glucose transporters like GLUT2[2-4] or SGLT2[5], and/or renal dysfunction, glucose is excreted in urine. Enhancing glycosuria, through SGLT2 inhibition[6-8], is now used to reduce blood glucose levels in humans with diabetes. However, the glucose loss is offset by a compensatory increase in endogenous glucose production, consequently compromising the efficacy of this strategy in treatment of diabetes[9, 10]. We recently produced renal Glut2 knockout mice[3], which show massive glycosuria and are protected from diabetes and diet-induced obesity. Despite the loss of glucose in urine, renal Glut2 knockout mice maintain normal fasting blood glucose levels. Collectively, these observations indicate the presence of a fundamental mechanism involved in sensing glucose loss and triggering a compensatory increase in endogenous glucose production in rodents and humans. This mechanism is also aligned with the established survival strategies that reduce utilization and increase production of glucose (energy) during a flight or fight response against a threat[11, 12].

Understanding how the kidney senses glucose loss and consequently activates compensatory pathways to make up for this loss will enhance our knowledge about integrative mechanisms that regulate glucose homeostasis. Answering such research questions may also be useful in medicine to increase efficiency of drugs such as SGLT2 inhibitors that are used to reduce blood glucose levels in diabetes by elevating glycosuria.

In this study, we used renal Glut2 knockout mice to identify mechanisms involved in activating a compensatory increase in glucose production in response to glucose loss in urine. We determined whether neural connections between the kidneys and hypothalamus can explain the compensatory increase in glucose production. In addition, we performed plasma proteomics analyses to identify secreted proteins or endocrine factors that may provide insights into pathways involved in conserving energy and producing endogenous glucose in response to elevated glycosuria in renal Glut2 knockout mice.

Methods

Mouse husbandry

Earlier version: All animal procedures were approved by the Institutional Animal Care and Use Committee at the University of Rochester, Yale University, or University of Iowa, and were performed according to the US Public Health Service guidelines for the humane care and use of experimental animals. Mice were housed in ventilated cages under controlled temperature (∼23°C) and photoperiod (12 h light / dark cycle, lights on from 06:00 hours to 18:00 hours or 07:00 hours to 19:00 hours at Yale) conditions with free access to Hydropac water (Lab Products, USA) and regular laboratory chow (5010, LabDiet, USA or Global 2018, Harlan Teklad, USA). We used male and female renal Glut2 knockout[3] and their littermate control mice in this study. After genotyping, mice were randomly assigned to different experimental groups. Renal Glut2 deficiency was induced using tamoxifen as described previously[3].

Revised version: All animal procedures were approved by the Institutional Animal Care and Use Committee at the University of Rochester, Yale University, or University of Iowa, and were performed according to the US Public Health Service guidelines for the humane care and use of experimental animals. Mice were housed in ventilated cages under controlled temperature (∼23°C) and photoperiod (12 h light / dark cycle, lights on from 06:00 hours to 18:00 hours or 07:00 hours to 19:00 hours at Yale) conditions with free access to Hydropac water (Lab Products, USA) and regular laboratory chow (5010, LabDiet, USA or Global 2018, Harlan Teklad, USA). Although we included all male or all female mice in a given experiment here, we did not repeat all the experiments in both sexes because our previous study[3] demonstrated that male and female renal Glut2 knockout mice exhibited a similar phenotype compared to their corresponding control groups in the context of glucose homeostasis. The age and sex of mice used in each experiment in this study are mentioned in the figure legends. After genotyping, mice were randomly assigned to different experimental groups. Renal Glut2 deficiency was induced using tamoxifen (50 mg/kg first dissolved in one part 100% ethanol followed by addition of nine parts of sesame oil, i.p., T5648 and S3547, Sigma, USA) as described previously[3].

Glucose production and liver metabolomics

Earlier version: Mice underwent surgery under isoflurane anesthesia to place a catheter in the jugular vein, and were allowed to recover for 7 days prior to flux studies. They were fasted overnight (16 h) to deplete liver glycogen prior to study. Mice were administered a 3X primed-continuous infusion of [3-13C] sodium lactate (continuous infusion rate 20 µmol/kg/min) and 2H7 glucose (continuous infusion rate 2.2 µmol/kg/min) for 120 min, and euthanized with IV Euthasol. Glucose production from pyruvate was measured in vivo in the liver and kidney (both freeze-clamped in situ within 30 second of euthanasia) using mass isotopomer distribution analysis as described previously[13]. Briefly, we measured the whole-body ratio of pyruvate carboxylase flux to total glucose production by gas chromatography/mass spectrometry-mass spectrometry. We compared the whole-body gluconeogenesis from phosphoenolpyruvate (PEP) and all substrates upstream of PEP (i.e., pyruvate, lactate, amino acids) to that measured locally in the liver and kidney. As the whole-body fraction of gluconeogenesis from PEP represents an integrated contribution of both liver and kidney to whole-body glucose production from PEP, this allowed us to –calculate the fractional contribution of the kidney to whole-body glucose production.

Revised version: Mice underwent surgery under isoflurane anesthesia to place a catheter in the jugular vein, and were allowed to recover for 7 days prior to flux studies. They were fasted overnight (16 h) to deplete liver glycogen prior to study. Mice were administered a 3X primed-continuous infusion of [3-13C] sodium lactate (continuous infusion rate 20 μmol/kg/min) and 2H7 glucose (continuous infusion rate 2.2 μmol/kg/min) for 120 min, and euthanized with IV Euthasol. Glucose production from pyruvate was directly measured in vivo in the liver and kidney (both freeze-clamped in situ within 30 second of euthanasia) using mass isotopomer distribution analysis as described previously[13]. Briefly, we measured the whole-body ratio of pyruvate carboxylase flux to total glucose production by gas chromatography/mass spectrometry-mass spectrometry. We compared the whole-body gluconeogenesis from phosphoenolpyruvate (PEP) and all substrates upstream of PEP (i.e., pyruvate, lactate, amino acids) to that measured locally in the liver and kidney. As the whole-body fraction of gluconeogenesis from PEP represents an integrated contribution of both liver and kidney to whole-body glucose production from PEP, this allowed us to –calculate the fractional contribution of the kidney to whole-body glucose production.

For measuring liver metabolites involved in glycolysis, flash-frozen liver was homogenized to extract metabolites in 80% methanol, which was evaporated under liquid nitrogen followed by resuspension of metabolites in 50% methanol. Liquid Chromatography Tandem Mass Spectrometry (LC-MS/MS) analysis was performed with reverse-phase LC on a Synergi Fusion-RP column (Phenomenex, Torrence CA) at 35 °C. Samples were analyzed by single reaction monitoring on a Thermo Quantum triple-quadrupole mass spectrometer. Metabolites were identified against a library of validated standards based on retention time, intact mass, collision energy, and fragment masses. Data were collected and analyzed using Thermo XCaliber 4.0 software. Post-hoc data was analyzed using Metaboanalyst 5.0. Data were median-normalized (normalized to the median peak area per sample, which correlated well with the total protein concentration). 73 metabolites were measured, out of which 9 were significantly (p<0.05) different (1.5 fold-change threshold, Supplementary data table 1).

Measurement of genes and hormones involved in regulating hypothalamic-pituitary-adrenal axis

Earlier version: We measured the levels of Crh mRNA in the mouse hypothalamus using RNA fluorescence in situ hybridization (Crh probe, 316091, ACD) according to the manufacturer’s instructions. DAPI was used to stain the nucleus and verify the regions of interest in the mouse brain. Images were captured using Keyence fluorescence microscope BZ-X800. To further quantify the number of cells expressing Crh or DAPI, we used CellProfiler 4.4.1 (https://cellprofiler.org/releases).

Revised version: We measured the levels of Crh mRNA in the mouse paraventricular hypothalamus using RNA fluorescence in situ hybridization (Crh probe, 316091, ACD) according to the manufacturer’s instructions. DAPI was used to stain the nucleus and verify the regions of interest in the mouse brain. Images were captured using Keyence fluorescence microscope BZ-X800. To further quantify the number of cells expressing Crh or DAPI, we used CellProfiler 4.4.1 (https://cellprofiler.org/releases).

We used the following primers to measure gene expression using RT-qPCR: Crh: 5’-GGAATCTCAACAGAAGTCCCGC-3’ and 5’-CTGCAGCAACACGCGGAAAAAG-3’, Hrprt: 5′-AACAAAGTCTGGCCTGTATCC-3′ and 5′-CCCCAAAATGGTTAAGGTTGC-3′. All primers were used at a final concentration of 500 nmol/l. The relative quantity of each mRNA was calculated from standard curves and normalized to the internal control Hprt, and then normalized to the mean of corresponding controls.

Added in the revised version: To minimize a natural stress response, mice were acclimated to handling procedures at least three days before collecting their blood. For measuring adrenocorticotropic hormone and corticosterone, we collected mouse tail blood at 10:00 hours on the day of measurements by nicking the tip of the tail with a sterile razor blade and then using heparinized capillary tubes (22-362566, Thermo Fisher Scientific). For measuring plasma insulin, mice were fasted for 6h (08:00 – 14:00 hours) before collecting their tail blood. The blood was centrifuged for 10 minutes at 2,000 x g and 40C. Plasma levels of adrenocorticotropic hormone (ab263880, Abcam), corticosterone (80556, Crystal Chem), and insulin (90080, Crystal Chem) were measured using ELISA per the manufacturers’ instructions.

Removed in the revised version: Plasma levels of adrenocorticotropic hormone (ab263880, Abcam) and corticosterone (80556, Crystal Chem) were measured using ELISA per the manufacturers’ instructions.

Ablation of afferent renal nerve

After anesthetizing mice using 1 – 5% isoflurane, a dorsal midline incision was made to access both the kidneys. We selectively ablated afferent renal nerves using capsaicin. Gauze, soaked in capsaicin (35 mM in 10% ethanol and 90% saline solution), was wrapped around the renal blood vessels for 15 min. Other surrounding tissues were covered with parafilm to avoid their exposure to capsaicin. After the surgery, muscle layers and skin were closed with nylon sutures. This procedure ablates afferent renal nerves without affecting the blood vessels or efferent renal nerves[14]. Control mice underwent the same procedure without capsaicin (sham surgery using 10% ethanol and 90% saline solution). The mice were allowed to recover for a week before determining the effects of the afferent denervation on blood glucose and other parameters.

Measurement of renal nerve activity, blood pressure, and heart rate

Earlier version: The mouse renal nerve activity was measured using multifiber recording as described previously [15]. After anesthesia, the nerve innervating the left kidney was identified, dissected free, and placed on a bipolar 36-gauge platinum–iridium electrode (A-M Systems; Carlsborg, WA). The electrode was connected to a high impedance probe (HIP-511; Grass Instruments Co., Quincy, MA), and the nerve signal was amplified 105 times with a Grass P5 AC pre-amplifier and filtered at low and high frequency cutoffs of 100 Hz and 1000 Hz, respectively. This nerve signal was directed to a speaker system and to an oscilloscope (54501A, Hewlett–Packard Co., Palo Alto, CA) for auditory and visual monitoring of the nerve activity. The signal was then directed to a resetting voltage integrator (B600C, University of Iowa Bioengineering) that sums the total voltage output in units of 1 V * sec before resetting to zero and counting the number of spikes per second. The final neurograms were continuously routed to a MacLab analogue–digital converter (8S, AD Instruments Castle Hill, New South Wales, Australia) for permanent recording and data analysis on a Macintosh computer. Renal sympathetic nerve activity was measured during 30 min. To measure the afferent renal nerve activity, the nerve was cut distally and the signal was recorded for 30 min. Nerve activity was corrected for post-mortem background activity to eliminate background electrical noise in the measurements.

Revised version: The mouse renal nerve activity was measured using multifiber recording as described previously [15]. After anesthesia, the nerve innervating the left kidney was identified, dissected free, and placed on a bipolar 36-gauge platinum–iridium electrode (A-M Systems; Carlsborg, WA). The electrode was connected to a high impedance probe (HIP-511; Grass Instruments Co., Quincy, MA), and the nerve signal was amplified 105 times with a Grass P5 AC pre-amplifier and filtered at low and high frequency cutoffs of 100 Hz and 1000 Hz, respectively. This nerve signal was directed to a speaker system and to an oscilloscope (54501A, Hewlett–Packard Co., Palo Alto, CA) for auditory and visual monitoring of the nerve activity. The signal was then directed to a resetting voltage integrator (B600C, University of Iowa Bioengineering) that sums the total voltage output in units of 1 V * sec before resetting to zero and counting the number of spikes per second. The final neurograms were continuously routed to a MacLab analogue–digital converter (8S, AD Instruments Castle Hill, New South Wales, Australia) for permanent recording and data analysis on a Macintosh computer. Renal sympathetic nerve activity was measured during 30 min. To measure the afferent renal nerve activity, the nerve was cut distally and the signal was recorded for 30 min. Nerve activity was corrected for post-mortem background activity to eliminate background electrical noise in the measurements. At the end of SNA recording, each mouse was euthanized with an overdose of anesthetic. Any remaining nerve activity after death was considered as background noise and subtracted from the SNA measurements In addition to the nerve activity, arterial pressure and heart rate were measured using a tapered MRE-040 tubing inserted into the left carotid artery and the other end connected to a transducer (BP-100; iWorks System, Inc., Dover, NH) that led to an

Earlier version: In addition to the nerve activity, arterial pressure and heart rate were measured using a tapered MRE-040 tubing inserted into the left carotid artery and the other end connected to a transducer (BP-100; iWorks System, Inc., Dover, NH) that led to an ETH-250 Bridge/Bio Amplifier (CB Sciences; Milford, MA). Core body temperature of the mouse was measured using a rectal probe (YSI 4000A Precision Thermometer; Yellow Springs, OH) and maintained constant at 37.5 °C using a custom-made heated surgical platform.

Revised version: ETH-250 Bridge/Bio Amplifier (CB Sciences; Milford, MA). Core body temperature of the mouse was measured using a rectal probe (YSI 4000A Precision Thermometer; Yellow Springs, OH) and maintained constant at 37.5 °C using a custom-made heated surgical platform.

Two-dimensional difference gel electrophoresis (2D-DIGE)

Plasma proteins from control and renal Glut2 knockout mice were labeled with Cy3 and Cy5 dyes respectively. We also used an internal standard (labeled with Cy2 dye) containing equal amounts of protein of each plasma sample as a quality control. Using 2D gel electrophoresis, we separated plasma proteins (30 μg in 2D sample buffer) on a gel by both isoelectric point and molecular weight. We then imaged the gel using Typhoon scanner. DeCyder software was used to identify protein spots of interest (downregulated or upregulated by at least 25%), which were automatically picked from the 2D gel with Ettan Spot Picker (Amersham BioSciences). The proteins were identified by MALDI TOF/TOF mass spectrometry (AB SCIEX TOF/TOF™ 5800 System) using Swiss-Prot database.

Statistics

Data are shown as mean□± s.e.m. Results were analyzed by two-tailed Student’s unpaired t-test or two-way ANOVA followed by a Tukey’s post hoc multiple comparison test when appropriate. All analyses were performed using Prism version 8.0.1 (GraphPad, USA) and differences were considered statistically significant at p<0.05.

Results and Discussion

Increased hepatic and renal glucose production in renal Glut2 knockout mice

Renal Glut2 knockout mice exhibit normal fasting blood glucose levels despite massive glycosuria[3]. This observation suggests a compensatory increase in endogenous glucose production, similar to that observed in humans following SGLT2 inhibition[9, 10, 16]. Therefore, we measured hepatic and renal glucose production in renal Glut2 knockout and their littermate control mice. We observed that both tissues contribute to the compensatory increase in glucose production in renal Glut2 knockout mice (Figure 1a-c). This was accompanied by a decrease in hepatic glucose 6-phosphate, glucose 1-phosphate, and fructose 6-phosphate (Figure 1d-f), suggesting a decline in glucose metabolism in favor of increasing glucose production. List of all measured metabolites is included in Supplementary data table 1.

Renal Glut2 knockout mice exhibit increased glucose production. In vivo increase in total (a), hepatic (b), and renal (c) glucose production through gluconeogenesis with pyruvate as a substrate in 28 weeks old male renal Glut2 knockout mice. Decreased hepatic glucose 6-phosphate (d), glucose 1-phosphate (e), and fructose 6-phosphate (f) in renal Glut2 knockout mice. *p<0.05, ***p<0.001, unpaired 2-tailed Student’s t-test. Data are presented as mean ±sem.

Renal Glut2 knockout mice have an activated hypothalamic-pituitary-adrenal (HPA) axis

The HPA axis is a major stress response system involved in enhancing glucose production. Therefore, we measured plasma corticosterone and adrenocorticotropin hormone coupled with gene expression of hypothalamic corticotropic releasing hormone (Crh). We observed an increase in all these factors (Figure 2) that comprise the HPA axis. These findings suggest that the HPA axis may contribute to the compensatory increase in glucose production in renal Glut2 knockout mice.

Earlier version: Enhanced activity of hypothalamic-pituitary-adrenal axis in renal Glut2 knockout mice. Representative images from fluorescence RNA in situ hybridization showing an increase in expression of corticotropin releasing hormone (Crh) in the paraventricular nucleus of the hypothalamus in 28 weeks old male renal Glut2 knockout mice (a). Scale, 100 µm. For the quantification shown next to the images, 4 sections per mouse and 3 areas of interest per section were analyzed in 4 mice. qRT-PCR analysis showing an increase in hypothalamic Crh (b), data from ELISA demonstrating an increase in plasma adrenocorticotropic hormone (ACTH) (c) and corticosterone (d) in 12 weeks old male renal Glut2 knockout mice. *p<0.05, **p<0.01, ***p<0.001, unpaired 2-tailed Student’s t-test. Data are presented as mean ±sem.

Revised version: Enhanced activity of hypothalamic-pituitary-adrenal axis in renal Glut2 knockout mice. Representative images from fluorescence RNA in situ hybridization showing an increase in expression of corticotropin releasing hormone (Crh) in the paraventricular nucleus of the hypothalamus (which is identified here using a white oval shape) in 28 weeks old male renal Glut2 knockout mice (a). Scale, 100 μm. For the quantification shown next to the images, 4 sections per mouse and 3 areas of interest per section were analyzed in 4 mice. qRT-PCR analysis showing an increase in hypothalamic Crh (b), data from ELISA demonstrating an increase in plasma adrenocorticotropic hormone (ACTH) (c) and corticosterone (d) in 12 weeks old male renal Glut2 knockout mice. *p<0.05, **p<0.01, ***p<0.001, unpaired 2-tailed Student’s t-test. Data are presented as mean ±sem.

Afferent renal nerves contribute to enhancing glucose production in renal Glut2 knockout mice

Earlier version: To determine the role of afferent renal nerves in triggering a compensatory glucose production in response to glycosuria, we used capsaicin to selectively suppress afferent renal nerve activity in renal Glut2 knockout mice. Afferent denervation was confirmed by measuring renal pelvic CGRP (pg/mg protein) at the end of this study (Sham: 86.3 ±6.4 (control mice), 93.7 ±8.2 (renal Glut2 knockout mice); Capsaicin: 4.7 ±0.06 (control mice), 6.4 ±0.7 (renal Glut2 knockout mice), using ELISA kit, Cayman Chemical, 589001). Mice were allowed to recover for a week before we fasted them to measure their blood glucose levels. Because reinnervation may occur following the denervation, we completed this study within 14 days after the denervation procedure. We observed that fasting and fed (random) blood glucose levels were about 50% decreased in renal Glut2 knockout mice compared to their control littermates (Figure 3a,b) after the afferent renal denervation. In addition, the denervation reversed the activation of the HPA axis in renal Glut2 knockout mice (Figure 3c,d). Fasting plasma insulin levels were not changed (Figure 3e) after the denervation. These observations suggest that renal nerves contribute to the compensatory increase in glucose production.

Revised version: To determine the role of afferent renal nerves in triggering a compensatory glucose production in response to glycosuria, we used capsaicin to selectively suppress afferent renal nerve activity in renal Glut2 knockout mice. Afferent denervation was confirmed by measuring renal pelvic CGRP (pg/mg protein) at the end of this study (Sham: 86.3 ±6.4 (control mice), 93.7 ±8.2 (renal Glut2 knockout mice); Capsaicin: 4.7 ±0.06 (control mice), 6.4 ±0.7 (renal Glut2 knockout mice), using ELISA kit, Cayman Chemical, 589001). Mice were allowed to recover for a week before we fasted them to measure their blood glucose levels. Because reinnervation may occur following the denervation, we completed this study within 14 days after the denervation procedure. We observed that fasting and fed (random) blood glucose levels were about 50% decreased in renal Glut2 knockout mice compared to their control littermates (Figure 3a,b) after the afferent renal denervation. In addition, the denervation reversed the activation of the HPA axis in renal Glut2 knockout mice (Figure 3c,d). Fasting plasma insulin levels were not changed (Figure 3e) after the denervation. These observations suggest that renal nerves partly contribute to the compensatory increase in glucose production.

Effects of afferent renal denervation on blood glucose levels and hypothalamic-pituitary-adrenal axis in renal Glut2 knockout mice. Afferent renal denervation decreases fed (random) and fasting (overnight, 6:00 pm – 9:00 am) blood glucose levels (a,b), restores expression of hypothalamic corticotropin releasing hormone (Crh) (c) measured using RT-qPCR, and plasma corticosterone (d) without affecting plasma insulin levels (e) in 30 weeks old female renal Glut2 knockout mice. **p<0.01, ***p<0.001, two-way ANOVA followed by a Tukey’s post hoc multiple comparison test. Data are presented as mean ±sem.

Earlier version: Although the afferent renal denervation attenuated the compensatory increase in glucose production in renal Glut2 knockout mice, there was no change in their total or afferent renal nerve activity (Supplementary Figure 1a,b). The Glut2 knockout mice also had similar blood pressure, heart rate coupled with similar weights of brown and white adipose tissue, kidney, and liver (Supplementary Figure 1c-j). These findings suggest the role of integrative neuroendocrine mechanisms independently of changes in nerve activity, fat, kidney, or liver mass, in triggering a compensatory increase in glucose production in renal Glut2 knockout mice. Overall, the results indicate that afferent renal nerves contribute to a compensatory increase in glucose production to make up for glucose loss in renal Glut2 knockout mice. These findings support data from humans[16] that renal nerves may be involved in mediating a compensatory increase in glucose production in response to elevated glycosuria by SGLT2 inhibition.

Revised version: Although the afferent renal denervation procedure attenuated the compensatory increase in glucose production in renal Glut2 knockout mice, unexpectedly there was no change in their total or afferent renal nerve activity (Supplementary Figure 1a,b). These findings suggest that neuroendocrine mechanisms independently of changes in renal nerve activity may be involved in increasing glucose production in response to elevated glycosuria in renal Glut2 knockout mice. In addition, the Glut2 knockout mice had similar blood pressure, heart rate coupled with similar weights of brown and white adipose tissue, kidney, and liver (Supplementary Figure 1c-j). Altogether, the results indicate that afferent renal nerves partly contribute to the compensatory increase in glucose production to make up for glucose loss in renal Glut2 knockout mice. These findings support the data from humans[16] that renal nerves may be involved in mediating a compensatory increase in glucose production in response to elevated glycosuria by SGLT2 inhibition.

Changes in circulating acute phase proteins in renal Glut2 knockout mice

To further investigate underlying mechanisms compensating for glucose loss in Glut2 knockout mice, we measured levels of secreted proteins that may explain the increase in endogenous glucose production. We performed 2D-DIGE followed by proteomic analyses with plasma samples collected from renal Glut2 knockout mice and their littermate controls. We observed that acute phase proteins like mannose binding lectin, albumin, haptoglobin, ferritin were either upregulated or downregulated in renal Glut2 knockout mice (Figure 4 and Supplementary table 2). Usually levels of these proteins are changed in response to threat, infection, and/or injury to activate body’s defense systems[17]. Similarly, major urinary proteins (MUP) like MUP18, which belong to the lipocalin protein family, were elevated in renal Glut2 knockout mice (Figure 4 and Supplementary table 2). These proteins are involved in regulating inter-organ chemical signaling, pheromonal communication, and energy metabolism[18]. We also observed that secreted glutathione peroxidase 3 (Gpx3) was the most downregulated protein in plasma in renal Glut2 knockout mice (Figure 4 and Supplementary table 2). Gpx3 is predominantly expressed in the kidneys and is an anti-oxidant[19]. We validated these findings using qRT-PCR. Renal Gpx3 gene expression was about 50% lower (Supplementary Figure 2a) in renal Glut2 knockout mice compared to their littermate controls. Moreover, renal - but not hepatic - Mup18 gene expression was higher (Supplementary Figure 2b,c) in renal Glut2 knockout mice relative to their controls, suggesting that the kidneys were responsible for increasing plasma MUP18 in the knockout mice. Altogether, these results suggest that multiple pathways - including local (renal) stress and general (HPA axis) stress response combined with the acute phase proteins - compensate for glucose loss and defend against a perceived biological threat in renal Glut2 knockout mice.

Changes in levels of plasma proteins in renal Glut2 knockout mice. Representative image of two-dimensional difference gel electrophoresis with numbered protein spots of interest is shown in (a). Internal standard was prepared using equal amounts of protein of each plasma sample as a quality control (Cy2 labeled, pseudo blue), plasma proteins from control group were labeled using Cy3 dye (shown in pseudo green), and plasma proteins from renal Glut2 knockout mice were labeled using Cy5 dye (shown in pseudo red). The identified proteins and their fold change in 28 weeks old male renal Glut2 knockout mice compared to the control group is shown in (b). The number on each bar graph in (b) represents the corresponding protein spot on the gel shown in (a). MBL1, mannose binding lectin 1; Gpx3, glutathione peroxidase 3; MUP18, major urinary protein 18.

Earlier version: This study has limitations. Higher urine volume (polyuria) observed in renal Glut2 knockout mice may contribute to increase in glucose production to some extent, however that contribution would likely be minor[20, 21] based on previous studies about the role of diuretics (which increase urine volume) in glucose production. The extent to which the changes in circulating secretory proteins observed in renal Glut2 knockout mice contribute to their compensatory increase in glucose production is unknown. Further studies are necessary to address the role of these secretory proteins in glucose and energy homeostasis. It is possible that afferent renal denervation in the present study attenuated only hepatic glucose production through the hypothalamus without affecting the compensatory increase in renal (local) glucose production. If validated, this would explain the source of remaining endogenous glucose after the denervation in renal Glut2 knockout mice.

Revised version: This study has limitations. Higher urine volume (polyuria) observed in renal Glut2 knockout mice may contribute to an increase in glucose production to some extent, however that contribution would likely be minor based on previous studies[20, 21] about the role of diuretics (which increase urine volume) in glucose production. Additional research is necessary to determine the role of the altered secretory proteins – that were identified through the plasma proteomics analyses – in triggering an increase in glucose production in response to elevated glycosuria. It is possible that afferent renal denervation in the present study attenuated only hepatic glucose production through the hypothalamus axis without affecting the compensatory increase in renal (which is the local site of denervation in this study) glucose production. If validated, this would explain why the afferent denervation did not completely block the compensatory glucose production in renal Glut2 knockout mice.

In summary, we report that a kidney-hypothalamus axis contributes to triggering a compensatory increase in endogenous glucose production in response to elevated glycosuria in mice. Glucose loss in urine appears to be sensed as a biological threat and therefore activates body’s defense systems including some of the acute phase proteins. These findings may explain why SGLT2 inhibitors don’t achieve their full potential in lowering blood glucose levels during treatment of diabetes mellitus.

Supporting information